---
title       : Prework2-Cox/ Poisson Modeling
description : Cox/ Poisson Modeling.
---

## <<<New Exercise>>>

```yaml
type: NormalExercise
key: b8586f931c
lang: r
xp: 50
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```
---
